Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Belle"
COLLABORATORS = "No one"
```

---

```python
def even(x):
    if x % 2 == 0:
        return True
    return False

```

<!-- grade_id: cell-6a704af2d61c8245, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert(even(0) == True)
assert(even(1) == False)
### END HIDDEN TESTS
```