Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "BELLE"
COLLABORATORS = "No one else"
```

---

<!-- grade_id: cell-6a704af2d61c8245, points: 4, grade: true, solution: false -->
```python
def even(x):
    ### DEFINE THIS
    return


### BEGIN HIDDEN TESTS
assert(even(0) == True)
assert(even(1) == False)
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[2], line 7
      3     return
      6 ### BEGIN HIDDEN TESTS
----> 7 assert(even(0) == True)
      8 assert(even(1) == False)

AssertionError:
```